<!-- grade_id: cell-70b8f5108418b01c, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    ### BEGIN SOLUTION
    if n < 1:
    ### END SOLUTION
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]

```

<!-- grade_id: cell-50b75e17d4507ac9, points: 50, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

<!-- grade_id: cell-4330a30cceaeaf3e, points: 50, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
### END HIDDEN TESTS
```